## Assignment-1
##### Total number of points: 50
#### Due date: 09/16/2024 5 PM ET

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management.  

#### Q1[5 points]. Variable Assignment

- Assign the value of `var11` to `20`. 
- Set the value for `var12` to `5` times `var11` , and set the value of `var13` to `var12` times by `var11`. 
- Assign `var14` to cube value of `var13` (meaning `var13` power `3`). 

<!-- grade_id: q1-sol, points: 0, grade: false, solution: true -->
```python
#Assign var11; var12; var13; var14

# YOUR CODE HERE
var11 = 20
var12 = 5 * var11
var13 = var12 * var11
var14 = var13 ** 3
# print(var11, var12, var13, var14)
```

<!-- grade_id: 1, points: 3, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-1 (3 points)
assert var11 == 20
assert var12 == 100
assert var14 == 8000000000
```

<!-- grade_id: 2, points: 2, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-2 (2 points)
```

#### Q2-1 [5 points]: List

- Create an empty list `var21`
- Create a list `var22` with integers `1`, `2`, `3` (all integers).
- Create a nested list `var23` with `var21` and `var22`
- Create a list `var24` with only `4` (integer) repeated 3 times.
- Create a list `var25` with your firstName and lastName as two separate strings -- for example I will build this as ['lydia', 'manikonda']


<!-- grade_id: q21-sol, points: 0, grade: false, solution: true -->
```python
#Need to create lists var21, var22, var23, var24, var25 according to the instructions

# YOUR CODE HERE
var21 = []
var22 = [1, 2, 3]
var23 = [var21, var22]
var24 = [4] * 3
var25 = ["Justin", "Chen"]
# print(var21, var22, var23, var24, var25)
```

<!-- grade_id: 3, points: 3, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-3 (3 points)
assert len(var21) == 0
assert set(var22) == {1,2,3}
assert len(var23) == 2
assert len(set(var24)) == 1

```

<!-- grade_id: 4, points: 2, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-4 (2 points)
```

#### Q2-2 [5 points]: 

###### Note that all are integers in this list. Please follow the exact order as  listed here below from first line to the last line
- Create a list `var21` with values in an order -- 10, -10, 20, -20 
- Append a new element `30` to list `var21`  
- Insert a new element `-30` at position `2`
- Insert a new element `15` at position `20` (Yes it is 20 -- check what happens)
- Create a new list `var22` rebinded as a new list with values only from `var21`. 
- Insert a new element `-10` at the end of the list `var21`. 

<!-- grade_id: q22-sol, points: 0, grade: false, solution: true -->
```python
#Need to create lists var21 and var22 according to the instructions above

# YOUR CODE HERE
var21 = [10, -10, 20, -20]
var21.append(30)
var21.insert(2, -30)
var21.insert(20, 15)
var22 = var21.copy()
var21.append(-10)
# print(var21, var22)
```

<!-- grade_id: 5, points: 2, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-5 (2 points)
assert len(var21)==8
assert var21[2]==-30
assert var21[7]==-10
assert var22[0] == 10

```

<!-- grade_id: 6, points: 3, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-6 (3 points)
```

#### Q3 [10 points] Dictionary

###### Create a dictionary `var31` with three keys 1, 2, 3 and their corresponding values as 2, 4, 6; Create another dictionary `var32` with two keys 4, 5 and their corresponding values as 8, 10 respectively. Merge these two dictionaries `var31` and `var32` to create a new dictionary `var33`. 

<!-- grade_id: q31-sol, points: 0, grade: false, solution: true -->
```python
#Creating the dictionaries var31, var32, var33

# YOUR CODE HERE
var31 = {
  1: 2,
  2: 4,
  3: 6,
}
var32 = {
  4: 8,
  5: 10
}
var33 = var31 | var32
# print(var33)
```

<!-- grade_id: 7, points: 3, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-7 (3 points)
assert len(var31)==3
assert len(var32)==2
assert var33[4]==8

```

<!-- grade_id: 8, points: 2, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-8 (2 points)
```

#### Q4-1 [5 points]. Set

- Create a set `S1` with values 1, 2, 3 in the same order as stated. 
- Assign the variable `Slen` as the length of this set `S1` 
- Create another set `S2` with values 4, 5, 6 in the same order as stated. 
- Create a new set `S3` by merging sets `S1` and `S2`

<!-- grade_id: q41-sol, points: 0, grade: false, solution: true -->
```python
#Please follow the above instructions to create S1, S2, S3

# YOUR CODE HERE
S1 = set([1, 2, 3])
Slen = len(S1)
S2 = set([4, 5, 6])
S3 = S1.union(S2)
# print(S1, S2, S3)
```

Output:
```
{1, 2, 3} {4, 5, 6} {1, 2, 3, 4, 5, 6}
```

<!-- grade_id: 9, points: 3, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-9 (3 points)
assert Slen==3
assert 1 in S1
assert S3=={1, 2, 3, 4, 5, 6}

```

<!-- grade_id: 10, points: 2, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-10 (2 points)
```

#### Q4-2 [5 points]

- Create a set `S4` with values 1, 2, 3, 4, 5 passed as strings 
- Add another string `'6'` to `S4`
- Convert each element in S4 which is a string to integer and save this new set as `S5`

<!-- grade_id: q42-sol, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
S4 = set(["1", "2", "3", "4", "5"])
S4.add("6")
S5 = set(int(x) for x in S4)
# print(S4, S5)
```

<!-- grade_id: 11, points: 5, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-11 (5 points)
assert set(S4)=={'1', '2', '3', '4', '5','6'}
assert S5=={1, 2, 3, 4, 5, 6}
```

#### Q5 [5 points]. Function

- Write a function `ReturnList` that takes a number `num` (where `num` is >0) and returns a list of numbers from 0 to `num` (inclusive). If number `num`<=0 return `[-1]`. 

<!-- grade_id: q5-sol, points: 0, grade: false, solution: true -->
```python

def ReturnList(num):

    # YOUR CODE HERE
    return [-1] if num <= 0 else list(range(0, num + 1))
ReturnList(4)
```

<!-- grade_id: 12, points: 3, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-12 (3 points)
assert ReturnList(3)==[0, 1, 2, 3]
assert ReturnList(4)==[0, 1, 2, 3, 4]
assert ReturnList(6)==[0, 1, 2, 3, 4, 5, 6]

```

<!-- grade_id: 13, points: 2, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-13 (2 points)
```

#### Q6 [5 points]. Lambda Function 

- Write a lambda function `CountVowels` that takes one string `str1` as input. This function should count and return an integer value representing the number of vowels in this string and assign this output to variable `var61`. 

- Vowels are {'a', 'e', 'i','o,'u'}



<!-- grade_id: q6-sol, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
CountVowels = lambda str1: sum(1 for char in str1 if char in 'aeiou')
# def CountVowels(str1):
#   return sum(1 for char in str1 if char in 'aeiou')
# print(CountVowels("hello world"))

```

Output:
```
3
```

<!-- grade_id: 14, points: 3, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-14 (3 points)
assert CountVowels('machine')==3
assert CountVowels('machine learning')==6

```

<!-- grade_id: 15, points: 2, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-15 (2 points)
```

#### Q7 [10 points]. Pattern Generation
- Write a function `GenPattern` that takes one integer as a parameter `limit` and generates a pattern with two characters as shown in the examples here.
![When limit = 7](https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/patternQuestion2.PNG)

![When limit = 17](https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/patternQuestion.PNG)

- You return this pattern as a nested list from the function. Assume that `limit` is always an odd number (limit%2 ==1). 

<!-- grade_id: q7-sol, points: 0, grade: false, solution: true -->
```python
def GenPattern(limit):
    # YOUR CODE HERE
    res = []
    # top half of the diamond
    for i in range(1, limit + 1, 2):
        res.append(' ' * ((limit - i) // 2) + '*' * i + ' ' * ((limit - i) // 2))
    # bottom half of the diamond
    for i in range(limit - 2, 0, -2):
        res.append(' ' * ((limit - i) // 2) + '*' * i + ' ' * ((limit - i) // 2))
    return res
print(GenPattern(9))
```

Output:
```
['    *    ', '   ***   ', '  *****  ', ' ******* ', '*********', ' ******* ', '  *****  ', '   ***   ', '    *    ']
```

<!-- grade_id: 16, points: 5, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-16 (5 points)
assert len(GenPattern(7))==7
assert GenPattern(7)[0].count(' ')==6
assert GenPattern(7)[3][0]=='*'
assert GenPattern(17)[8][0]=='*'
assert GenPattern(17)[8].count('*')==17
```

<!-- grade_id: 17, points: 5, grade: true, solution: false -->
```python
#Do not change this cell.
#Test cell-17 (5 points)
```

```python

```