## Assignment-3
##### Total number of points: 50
#### Due date: 10/7/2024 11:59 PM

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding.
- Failure to use proper letter casing for variable names will lead to loss of points. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management.  

#### We will be using the popular breast cancer wisconsin dataset for this homework (more details here: [url](https://scikit-learn.org/stable/modules/generated/sklearn.datasets.load_breast_cancer.html))

#### Q1. Metadata and building the feature set and target variable. 
1. Load the data separately as `X` and `y_true` corresponding to features and class label or target variable respectively using the `load_breast_cancer()` function. When you build the feature set `X` make sure `X` is a numpy's ndarray. More details about `load_breast_cancer()` function can be found here in the Examples section: [url](https://scikit-learn.org/stable/modules/generated/sklearn.datasets.load_breast_cancer.html)  <br>

2. Save all the `feature names` as a list of strings named as `colnames`.<br>

3. Convert `X` from numpy ndarray to a pandas dataframe (name this as `df_X`) that contains both data and the corresponding feature names.<br>

4. Perform standardization on all the features in `df_X` and save the standardized data as a dataframe named `df_X_std`.



<!-- grade_id: Q1-Sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_breast_cancer
import math

# YOUR CODE HERE
# Load the breast cancer dataset
data = load_breast_cancer()
X = data.data
y_true = data.target
colnames = data.feature_names
df_X = pd.DataFrame(X, columns=colnames)
df_X_std = df_X.std()
# print(df_X_std)
print(y_true)


```

Output:
```
[0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 1 1 1 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0
 1 0 0 0 0 0 0 0 0 1 0 1 1 1 1 1 0 0 1 0 0 1 1 1 1 0 1 0 0 1 1 1 1 0 1 0 0
 1 0 1 0 0 1 1 1 0 0 1 0 0 0 1 1 1 0 1 1 0 0 1 1 1 0 0 1 1 1 1 0 1 1 0 1 1
 1 1 1 1 1 1 0 0 0 1 0 0 1 1 1 0 0 1 0 1 0 0 1 0 0 1 1 0 1 1 0 1 1 1 1 0 1
 1 1 1 1 1 1 1 1 0 1 1 1 1 0 0 1 0 1 1 0 0 1 1 0 0 1 1 1 1 0 1 1 0 0 0 1 0
 1 0 1 1 1 0 1 1 0 0 1 0 0 0 0 1 0 0 0 1 0 1 0 1 1 0 1 0 0 0 0 1 1 0 0 1 1
 1 0 1 1 1 1 1 0 0 1 1 0 1 1 0 0 1 0 1 1 1 1 0 1 1 1 1 1 0 1 0 0 0 0 0 0 0
 0 0 0 0 0 0 0 1 1 1 1 1 1 0 1 0 1 1 0 1 1 0 1 0 0 1 1 1 1 1 1 1 1 1 1 1 1
 1 0 1 1 0 1 0 1 1 1 1 1 1 1 1 1 1 1 1 1 1 0 1 1 1 0 1 0 1 1 1 1 0 0 0 1 1
 1 1 0 1 0 1 0 1 1 1 0 1 1 1 1 1 1 1 0 0 0 1 1 1 1 1 1 1 1 1 1 1 0 0 1 0 0
 0 1 0 0 1 1 1 1 1 0 1 1 1 1 1 0 1 1 1 0 1 1 0 0 1 1 1 1 1 1 0 1 1 1 1 1 1
 1 0 1 1 1 1 1 0 1 1 0 1 1 1 1 1 1 1 1 1 1 1 1 0 1 0 0 1 0 1 1 1 1 1 0 1 1
 0 1 0 1 1 0 1 0 1 1 1 1 1 1 1 1 0 0 1 1 1 1 1 1 0 1 1 1 1 1 1 1 1 1 1 0 1
 1 1 1 1 1 1 0 1 0 1 1 0 
```

<!-- grade_id: 1, points: 5, grade: true, solution: false -->
```python
#Test cell-1 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert isinstance(X,np.ndarray)==True
assert isinstance(y_true,np.ndarray)==True
assert len(X)==569
```

<!-- grade_id: 2, points: 5, grade: true, solution: false -->
```python
#Test cell-2 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert sum(y_true)==357
assert df_X.isnull().values.sum()==0
assert len(set(colnames))==30
assert 'symmetry error' in colnames
```

<!-- grade_id: 3, points: 5, grade: true, solution: false -->
```python
#Test cell-3 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q2. Sampling with replacement
1. Create a new dataframe `ndf` using `df_X` from the previous question by randomly sampling the data points with replacement. Sample size is 400. 

Please note that we are not doing sampling without replacement here which means, there is a possibility of repeated datapoints in `ndf`. 

<!-- grade_id: Q2-Sol, points: 0, grade: false, solution: true -->
```python
#creating a new dataframe ndf 
# YOUR CODE HERE
ndf = df_X.sample(n=400, replace=True)
print(ndf)
```

Output:
```
mean radius  mean texture  mean perimeter  mean area  mean smoothness  \
529       12.070         13.44           77.83      445.2          0.11000   
438       13.850         19.60           88.68      592.6          0.08684   
459        9.755         28.20           61.68      290.9          0.07984   
136       11.710         16.67           74.72      423.6          0.10510   
545       13.620         23.23           87.19      573.2          0.09246   
..           ...           ...             ...        ...              ...   
3         11.420         20.38           77.58      386.1          0.14250   
363       16.500         18.29          106.60      838.1          0.09686   
292       12.950         16.02           83.14      513.7          0.10050   
341        9.606         16.84           61.64      280.5          0.08481   
141       16.110         18.05          105.10      813.0          0.09721   

     mean compactness  mean concavity  mean concave points  mean sym
```

<!-- grade_id: 4, points: 5, grade: true, solution: false -->
```python
#Test cell-4 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q3. Robustness metrics. 
1. Lets focus on the target variable or the class label column `y_true`. 
2. Create two lists `y_pred` and `y_pred_prob` which are of the same length as `y_true`.
    - assume that `y_pred` is the list of predicted class labels from the classifier.
    - `y_pred_prob` is the list of class probabilities from the classifier.
    - Assume that all the values in `y_pred` are `1`s and `y_pred_prob` are all `0.4` which is a vector of probability estimates.
3. Use the predicted class labels `y_pred` to compute precision, recall, accuracy values and save them as `precision`, `recall`, and `accuracy` respectively.
4. Also compute the area under curve and save it as variable name `auc_score`.

Please make sure you have all the variables assigned with the values and use the same letter casing as used here. 

<!-- grade_id: Q3-sol, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import precision_score, recall_score, accuracy_score, roc_auc_score


# YOUR CODE HERE
# Create y_pred and y_pred_prob
y_pred = [1] * len(y_true)
y_pred_prob = [0.4] * len(y_true)

# Compute precision, recall, and accuracy
precision = precision_score(y_true, y_pred)
recall = recall_score(y_true, y_pred)
accuracy = accuracy_score(y_true, y_pred)

# Compute area under curve
auc_score = roc_auc_score(y_true, y_pred_prob)
```

<!-- grade_id: 5, points: 10, grade: true, solution: false -->
```python
#Test cell-5 (10 points)
#DO NOT MODIFY/DELETE THIS CELL
assert round(precision,2) == 0.63
assert round(recall,2) == 1.0
assert round(accuracy,2) == 0.63
```

<!-- grade_id: 6, points: 5, grade: true, solution: false -->
```python
#Test cell-6 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q4. Data preprocessing and dimensionality reduction using PCA. 

1. Consider the same dataset ([url](https://scikit-learn.org/stable/modules/generated/sklearn.datasets.load_breast_cancer.html)) to create `X` and `y` where, `X` corresponds to only the feature columns, and `y` corresponds to the class label column `y_true`. Please make sure you check the casing -- `X` and `y` are in uppercase and lowercase formats respectively. <br><br>

2. Standardize all the features (`X`) using StandardScaler function and save its output as `X_std`. <br><br>

3. Consider only the top-`k` Principal components that that will retain not more than 90% of the information present in the original dataset. In other words just choose the first `k` principal components that capture 90% of variance only. Please assign the integer variable `k` with the total number of these top principal components. <br><br>

4. Translate the data only into those `k` new feature (which are the principal components) space and save the transformed data as `X_pca`. <br><br>

<!-- grade_id: Q4-sol, points: 0, grade: false, solution: true -->
```python
#Perform standardization using StandardScaler function and perform PCA.

#Consider only the top-k Principal components that that will retain not more than 78% of the information present in the original dataset.

#In other words just choose the first k principal components that capture 78% of variance only.

#How many such PCs should be considered?

import numpy as np
from sklearn.preprocessing import StandardScaler
from sklearn.decomposition import PCA

# YOUR CODE HERE
# Standardize the features
scaler = StandardScaler()
X_std = scaler.fit_transform(X)

# Perform PCA to retain 90% of the variance
pca = PCA(n_components=6)
X_pca = pca.fit_transform(X_std)

# Number of principal components
k = X_pca.shape[1]
print(k)

```

Output:
```
6
```

<!-- grade_id: 7, points: 5, grade: true, solution: false -->
```python
#Test cell-7 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert len(X)==569
assert len(X[0])==30
assert len(X)==len(y_true)
```

<!-- grade_id: 8, points: 5, grade: true, solution: false -->
```python
#Test cell-8 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert round(np.mean(X_std),2)==0
assert round(np.std(X_std),2)==1
assert len(X_pca)==len(X)
assert k>0
```

<!-- grade_id: 9, points: 5, grade: true, solution: false -->
```python
#Test cell-9 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```