## Assignment-5
##### Total number of points: 45
#### Due date: 11/7/2024 11:59 PM

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding.
- Failure to use proper letter casing for variable names will lead to loss of points. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management.  

#### Q1 [10 points]. We have loaded the wine dataset (from sklearn.datasets) -- use a dataframe 'df' to perform the following tasks:

###### 1. **Without using train_test_split() function** 
- Use the last 78 rows (i.e., index 100,101,....177) for training. Where, the corresponding features are loaded as `train_X` and corresponding labels as `train_y`.
- Use the remaining rows as testing data -- `test_X` and `test_y` for the features and labels respectively.
- Note that `train_X`, `train_y`,  `test_X` and `test_y` are dataframes.
- Fit a logistic regression with solver=`newton-cg`, C=1e5, multi_class=`multinomial`; then use it for predicting testing data as shown in the lecture notebook. 

###### 2. Measure the accuracy `acc_split1` using the predicted labels with groundtruth test labels `test_y`.
- Round the `acc_split1` to 2 values after the decimal point. 
- Hint -- Check out the libraries give


<!-- grade_id: Q1-Sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_wine
from sklearn import datasets
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score

#Load the dataset without class labels and save it as a dataframe
wine = load_wine()
df = pd.DataFrame(wine.data, columns=wine.feature_names)
#class label -- wine.target
#print(len(df))

train_X = pd.DataFrame()
train_y = pd.DataFrame()
test_X = pd.DataFrame()
test_y = pd.DataFrame()

acc_split1=0

# YOUR CODE HERE
X, y = load_wine(as_frame=True, return_X_y=True)
train_X = X[100:]
train_y = y[100:]
test_X = X[:100]
test_y = y[:100]

# Fit the logistic regression model
model = LogisticRegression(solver='newton-cg', C=1e5, multi_class='multinomial')
model.fit(train_X, train_y)

# Predict the testing data
pred_y = model.predict(test_X)

# Measure the accuracy
acc_split1 = round(accuracy_score(test_y, pred_y), 2)
acc_split1, train_X.columns, train_y.shape, test_X.shape, test_y.shape
```

Output:
```
c:\Users\Justin\anaconda3\envs\MGMT-4190\Lib\site-packages\sklearn\linear_model\_logistic.py:1237: FutureWarning: 'multi_class' was deprecated in version 1.5 and will be removed in 1.7. From then on, binary problems will be fit as proper binary  logistic regression models (as if multi_class='ovr' were set). Leave it to its default value to avoid this warning.
  warnings.warn(
(0.36,
 Index(['alcohol', 'malic_acid', 'ash', 'alcalinity_of_ash', 'magnesium',
        'total_phenols', 'flavanoids', 'nonflavanoid_phenols',
        'proanthocyanins', 'color_intensity', 'hue',
        'od280/od315_of_diluted_wines', 'proline'],
       dtype='object'),
 (78,),
 (100, 13),
 (100,))
```

<!-- grade_id: c1, points: 0, grade: true, solution: false -->
```python
#check for variable names
assert 'train_X' in globals()
assert 'train_y' in globals()
assert 'test_X' in globals()
assert 'test_y' in globals()
assert 'acc_split1' in globals()
```

<!-- grade_id: 1, points: 5, grade: true, solution: false -->
```python
#Test cell-1 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert len(train_X)==78
assert len(test_y)==100
assert acc_split1==0.36
```

<!-- grade_id: 2, points: 5, grade: true, solution: false -->
```python
#Test cell-2 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q2 [10 points].  Repeat the above exercise (Q1) but with different training and testing splits using the `wine` dataset.

###### 1. **Without using train_test_split() function** 
- Use the first 100 rows (i.e., index 0,1,....99) as training data -- `train_X` and `train_y` for the features and labels respectively.
- Use the last 78 rows (i.e., index 100,101,....177) for testing. Where, the corresponding features are loaded as `test_X` and corresponding labels as `test_y`.
- Note that `train_X`, `train_y`,  `test_X` and `test_y` are dataframes.
- Fit a logistic regression with solver=`newton-cg`, C=1e5, multi_class=`multinomial`; then use it for predicting testing data as shown in the lecture notebook. 

###### 2. Measure the accuracy `acc_split2` using the predicted labels with groundtruth test labels `test_y`.
- Round the `acc_split2` to 2 values after the decimal point. 
- Hint -- Check out the libraries give


<!-- grade_id: Q2-Sol, points: 0, grade: false, solution: true -->
```python

import pandas as pd
import numpy as np
from sklearn.datasets import load_wine
from sklearn import datasets
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score

#Load the dataset without class labels and save it as a dataframe
wine = load_wine()
df = pd.DataFrame(wine.data, columns=wine.feature_names)
#class label -- wine.target

train_X = pd.DataFrame()
train_y = pd.DataFrame()
test_X = pd.DataFrame()
test_y = pd.DataFrame()

acc_split2=0

# YOUR CODE HERE
X, y = load_wine(as_frame=True, return_X_y=True)
train_X = X[:100]
train_y = y[:100]
test_X = X[100:]
test_y = y[100:]

# Fit the logistic regression model
model = LogisticRegression(solver='newton-cg', C=1e5, multi_class='multinomial')
model.fit(train_X, train_y)

# Predict the testing data
pred_y = model.predict(test_X)

# Measure the accuracy
acc_split2 = round(accuracy_score(test_y, pred_y), 2)
acc_split2, train_X.columns, train_y.shape, test_X.shape, test_y.shape
```

Output:
```
c:\Users\Justin\anaconda3\envs\MGMT-4190\Lib\site-packages\sklearn\linear_model\_logistic.py:1237: FutureWarning: 'multi_class' was deprecated in version 1.5 and will be removed in 1.7. From then on, binary problems will be fit as proper binary  logistic regression models (as if multi_class='ovr' were set). Leave it to its default value to avoid this warning.
  warnings.warn(
(0.35,
 Index(['alcohol', 'malic_acid', 'ash', 'alcalinity_of_ash', 'magnesium',
        'total_phenols', 'flavanoids', 'nonflavanoid_phenols',
        'proanthocyanins', 'color_intensity', 'hue',
        'od280/od315_of_diluted_wines', 'proline'],
       dtype='object'),
 (100,),
 (78, 13),
 (78,))
```

<!-- grade_id: c2, points: 0, grade: true, solution: false -->
```python
#check for variable names
assert 'train_X' in globals()
assert 'train_y' in globals()
assert 'test_X' in globals()
assert 'test_y' in globals()
assert 'acc_split2' in globals()
```

<!-- grade_id: 3, points: 5, grade: true, solution: false -->
```python
#Test cell-3 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert len(train_X)==100
assert len(test_y)==78
assert acc_split2==0.35
```

<!-- grade_id: 4, points: 5, grade: true, solution: false -->
```python
#Test cell-4 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q3 [10 points]. Please refer to the data table shown in Lecture-14 (slide-33). 
- Compute entropy after splitting the original data using the attribute `credit_rating` and save it as variable `entropy_cr`. 
- Compute the information gain (use the variable `info_gain`) if we split the original data using `credit_rating` this information gain computation requires you to compute the total entropy before splitting the data.

<!-- grade_id: Q3-sol, points: 0, grade: false, solution: true -->
```python

import math
entropy_cr=-1 #this has to be updated using your computations
info_gain=-1 #this has to be updated

# YOUR CODE HERE
# Define the dataset
data = {
    'age': ['<=30', '<=30', '31...40', '>40', '>40', '>40', '31...40', '<=30', '<=30', '>40', '<=30', '31...40', '31...40', '>40'],
    'income': ['high', 'high', 'high', 'medium', 'low', 'low', 'low', 'medium', 'low', 'medium', 'medium', 'medium', 'high', 'medium'],
    'student': ['no', 'no', 'no', 'no', 'yes', 'yes', 'yes', 'no', 'yes', 'yes', 'yes', 'no', 'yes', 'no'],
    'credit_rating': ['fair', 'excellent', 'fair', 'fair', 'fair', 'excellent', 'excellent', 'fair', 'fair', 'fair', 'excellent', 'excellent', 'fair', 'excellent'],
    'buys_computer': ['no', 'no', 'yes', 'yes', 'yes', 'no', 'yes', 'no', 'yes', 'yes', 'yes', 'yes', 'yes', 'no']
}

df = pd.DataFrame(data)
df

# Function to calculate entropy
def entropy(labels):
  value, counts = np.unique(labels, return_counts=True)
  probs = counts / len(labels)
  return -np.sum(probs * np.log2(probs))

# Calculate total entropy before splitting
total_entropy = entropy(df['buys_computer'])

# Calculate entropy after splitting by credit_rating
entropy_cr = 0
for value in df['credit_rating'].unique():
  subset = df[df['credit_rating'] == value]
  entropy_cr += (len(subset) / len(df)) * entropy(subset['buys_computer'])

# Calculate information gain
info_gain = total_entropy - entropy_cr

entropy_cr, info_gain

```

Output:
```
(np.float64(0.8921589282623617), np.float64(0.04812703040826949))
```

<!-- grade_id: c3, points: 0, grade: true, solution: false -->
```python
#check for variable names
assert 'entropy_cr' in globals()
assert 'info_gain' in globals()
```

<!-- grade_id: 5, points: 5, grade: true, solution: false -->
```python
#Test cell-5 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert round(entropy_cr, 2)==0.89
```

<!-- grade_id: 6, points: 5, grade: true, solution: false -->
```python
#Test cell-6 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q4 [10 points]. Decision Tree Classifier  

###### 1. **Using train_test_split() function** 
- Create `train_X`, `train_y`,  `test_X` and `test_y` as dataframes corresponding to the training and testing datasets using `80:20` ratio for training and testing data split and with a `random_state` as `5`.
- Fit a Decision Tree Classifier `dt` with criterion as `entropy` and a `random_state` as `10` using training data; then use it for predicting testing data. 

###### 2. Measure the accuracy `acc_split3` using the predicted labels with groundtruth test labels `test_y`.
- Round the `acc_split3` to 2 values after the decimal point. 
- Hint -- Check out the libraries give


<!-- grade_id: Q4-sol, points: 0, grade: false, solution: true -->
```python

import pandas as pd
import numpy as np
from sklearn.datasets import load_wine
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score
from sklearn.tree import DecisionTreeClassifier

#Load the dataset without class labels and save it as a dataframe
wine = load_wine()
df = pd.DataFrame(wine.data, columns=wine.feature_names)
#class label -- wine.target

train_X = pd.DataFrame()
train_y = pd.DataFrame()
test_X = pd.DataFrame()
test_y = pd.DataFrame()

acc_split3 = 0
# YOUR CODE HERE
train_X, test_X, train_y, test_y = train_test_split(df, wine.target, test_size=0.2, random_state=5)

dt = DecisionTreeClassifier(criterion='entropy', random_state=10)
dt.fit(train_X, train_y)

# Predict the testing data
pred_y = dt.predict(test_X)

# Measure the accuracy
acc_split3 = round(accuracy_score(test_y, pred_y), 2)
acc_split3

```

Output:
```
0.81
```

<!-- grade_id: c4, points: 0, grade: true, solution: false -->
```python
#check for variable names
assert 'train_X' in globals()
assert 'train_y' in globals()
assert 'test_X' in globals()
assert 'test_y' in globals()
assert 'acc_split3' in globals()
assert 'dt' in globals()
```

<!-- grade_id: 7, points: 5, grade: true, solution: false -->
```python
#Test cell-7 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert len(train_X)==142
assert len(test_y)==36
assert dt.criterion=='entropy'
```

<!-- grade_id: 8, points: 5, grade: true, solution: false -->
```python
#Test cell-8 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q5 [5 points]. TF-IDF. 
- Given a corpus with three documents `D1`, `D2`, and `D3`: 
<ul>
<li>D1: The quick red fox jumps over the lazy dog.</li>
<li>D2: A quick red fox jumps over the lazy dog.</li>
<li>D3: The lazy dog watches as a quick red fox jumps over it.</li>
</ul>
- Convert all the documents into lowercase. 
- Remove all non-alphanumeric characters in all these documents.
- Then perform TF-IDF vectorization of these documents and save them as `TD1`, `TD2`, and `TD3` respectively.
- Compute the cosine similarity values and find the two documents with the closest similarity. Save this as a list variable `sim_docs` that contains the transformed names of these two documents to receive full points. For instance, if `D3` and `D1` are the most similar, `sim_docs` will be `['TD3', 'TD1']`.

<!-- grade_id: Q5-sol, points: 0, grade: false, solution: true -->
```python
#Find the two most similar documents using TF-IDF vectorization
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.metrics.pairwise import cosine_similarity
sim_docs=[]

# YOUR CODE HERE
D1 = 'The quick red fox jumps over the lazy dog.'
D2 = 'A quick red fox jumps over the lazy dog.'
D3 = 'The lazy dog watches as a quick red fox jumps over it.'
D1 = D1.lower().replace('.','')
D2 = D2.lower().replace('.','')
D3 = D3.lower().replace('.','')
tfidf_matrix = TfidfVectorizer().fit_transform([D1, D2, D3])
TD1, TD2, TD3 = tfidf_matrix.toarray()
# print(tfidf_matrix)
cosine_sim = cosine_similarity(tfidf_matrix)
print(cosine_sim)
sim_index = np.unravel_index(np.argmax(cosine_sim - np.eye(cosine_sim.shape[0])), cosine_sim.shape)
sim_docs = [f'TD{sim_index[0] + 1}', f'TD{sim_index[1] + 1}']
sim_docs
```

Output:
```
[[1.         0.95940322 0.66602258]
 [0.95940322 1.         0.69420507]
 [0.66602258 0.69420507 1.        ]]
['TD1', 'TD2']
```

<!-- grade_id: c5, points: 0, grade: true, solution: false -->
```python
#check for variable names
assert 'D1' in globals()
assert 'D2' in globals()
assert 'D3' in globals()
assert 'sim_docs' in globals()
```

<!-- grade_id: 9, points: 5, grade: true, solution: false -->
```python
#Test cell-9 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```