## Assignment-6
##### Total number of points: 40
#### Due date: 11/14/2024 11:59 PM

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding.
- Failure to use proper letter casing for variable names will lead to loss of points. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management.  

```python
#Please run this cell to load the dataset. 
url = 'https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/dataset_Facebook.csv'

import pandas as pd
import math
df  = pd.read_csv(url)
df = df.dropna()
df.columns
```

Output:
```
Index(['Page total likes', 'Type', 'Category', 'Post Month', 'Post Weekday',
       'Post Hour', 'Paid', 'Lifetime Post Total Reach',
       'Lifetime Post Total Impressions', 'Lifetime Engaged Users',
       'Lifetime Post Consumers', 'Lifetime Post Consumptions',
       'Lifetime Post Impressions by people who have liked your Page',
       'Lifetime Post reach by people who like your Page',
       'Lifetime People who have liked your Page and engaged with your post',
       'comment', 'like', 'share', 'Total Interactions'],
      dtype='object')
```

#### Part-1 [24 points]: We will leverage randomforest classifier to identify the best features that predict 'Total Interactions'.

#### 1. Create a new dataframe `df_sub2` that contains only these feature columns in `df` : 
- `Page total likes`, `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`, `Paid`,  `Total Interactions`

#### 2. Transform the categorical attribute `Type` in `df_sub2` to numerical attribute this way: 
- `Link`:1, `Photo`:2, `Status`:3, `Video`:4

#### 3. Perform Standardization (using the formula we learnt) only on `Page total likes` column in `df_sub2` 
- Please use `<dataframe>.mean()` `<dataframe>.std()` if you are using the mean and std values to manipulate the column. 
    
#### 4. Using `df_sub2` perform train_test_split operation to build training (`X_train`, `y_train`) and testing data (`X_test`, `y_test`).
- Use test_size=0.3, random_state=42 as the parameters for train_test_split() function. 
- Feature columns (X): `Page total likes`, `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`, `Paid`
- CLass label column (y): `Total Interactions`

#### 5. Train the randomforest classifier (initialized as variable `rf`) using these parameters: max_depth=3, random_state=0. 
- Using the trained model `rf` compute the impurity-based feature importances.
- Append the names of these top-3 features to list `impFeatures`. Please make sure you type the feature names exactly as in df_sub2 




<!-- grade_id: q1-sol, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestClassifier
from sklearn.model_selection import train_test_split
import numpy as np

# YOUR CODE HERE
# 1
df_sub2 = df[['Page total likes', 'Type', 'Category', 'Post Month', 'Post Weekday', 'Post Hour', 'Paid', 'Total Interactions']]

# 2
df_sub2['Type'] = df_sub2['Type'].map({'Link': 1, 'Photo': 2, 'Status': 3, 'Video': 4})

# 3
df_sub2['Page total likes'] = (df_sub2['Page total likes'] - df_sub2['Page total likes'].mean()) / df_sub2['Page total likes'].std()

# 4
X = df_sub2[['Page total likes', 'Type', 'Category', 'Post Month', 'Post Weekday', 'Post Hour', 'Paid']]
y = df_sub2['Total Interactions']
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=42)

# 5
rf = RandomForestClassifier(max_depth=3, random_state=0)
rf.fit(X_train, y_train)
importances = rf.feature_importances_
indices = np.argsort(importances)[-3:][::-1]
impFeatures = [X.columns[i] for i in indices]
```

Output:
```
C:\Users\Justin\AppData\Local\Temp\ipykernel_11884\1820953736.py:10: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  df_sub2['Type'] = df_sub2['Type'].map({'Link': 1, 'Photo': 2, 'Status': 3, 'Video': 4})
C:\Users\Justin\AppData\Local\Temp\ipykernel_11884\1820953736.py:13: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  df_sub2['Page total likes'] = (df_sub2['Page total likes'] - df_sub2['Page total likes'].mean()) / df_sub2['Page total likes'].std()
```

<!-- grade_id: t1, points: 0, grade: true, solution: false -->
```python
#check for variable names
assert 'df_sub2' in globals()
assert 'X_train' in globals()
assert 'impFeatures' in globals()
assert 'rf' in globals()
assert 'y_test' in globals()
```

<!-- grade_id: 1, points: 6, grade: true, solution: false -->
```python
#Test cell-1 (6 points)
#DO NOT MODIFY/DELETE THIS CELL
assert set(df_sub2.columns)=={'Post Hour', 'Total Interactions', 'Post Weekday', 'Category', 'Paid', 'Post Month', 'Page total likes', 'Type'}
assert len(df_sub2)==495
assert len(df_sub2.iloc[0,:])==8

```

<!-- grade_id: 2, points: 6, grade: true, solution: false -->
```python
#Test cell-2 (6 points)
#DO NOT MODIFY/DELETE THIS CELL
assert math.ceil(df_sub2['Page total likes'].mean())==1
assert len(X_train)==346
assert 'Post Hour' in impFeatures

```

<!-- grade_id: 3, points: 6, grade: true, solution: false -->
```python
#Test cell-3 (6 points)
#DO NOT MODIFY/DELETE THIS CELL
assert rf.n_classes_==230
assert len(impFeatures)==3
assert math.ceil(df_sub2['Page total likes'].std())==1
```

<!-- grade_id: 4, points: 6, grade: true, solution: false -->
```python
#Test cell-4 (6 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Part-2 [16 points]: We will now use neural networks to model a regression problem. 

#### 1. Create a new dataframe `df_sub4` that contains only these feature columns in `df` : 
- `Page total likes`, `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`, `Paid`,  `Total Interactions`

#### 2. Perform one-hot encoding on these features below in the dataframe `df_sub4` to create a new dataframe `df_OHE`
- `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`

#### 3. Perform normalization only on `Page total likes` and `Total Interactions` column in `df_OHE` using MinMaxScaler or minmax_scale 
- note that both these functions will output the same result

#### 4. Using `df_OHE` perform train_test_split operation to build training (`X_train`, `y_train`) and testing data (`X_test`, `y_test`).
- Use test_size=0.3, random_state=42 as the parameters for train_test_split() function. 
- Feature columns (X): Everything except `Total Interactions`
- CLass label column (y): `Total Interactions`

#### 5. Build a 3-layer neural network with `relu` as the 2 hidden layers' activation function and `sigmoid` as the final layer's activation function. 
- Ensure your `model` is `Sequential`.
- Use the same `adam` optimizer and `mean_squared_error` as the loss function. 
- `epochs`=150, `batch_size`=10
- since the class label is continuous, we use `metrics.MeanSquaredError()` as the metric (`from keras import metrics`) to measure the performance of our classifier captured as variable `mse` 

```python
#Please run this cell to install keras
!pip install keras
!pip install tensorflow
```

<!-- grade_id: q2-sol, points: 0, grade: false, solution: true -->
```python
from sklearn.preprocessing import MinMaxScaler
from sklearn.preprocessing import minmax_scale
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score
from keras.models import Sequential
from keras.layers import Dense
from keras import metrics

# YOUR CODE HERE
# 1
df_sub4 = df[['Page total likes', 'Type', 'Category', 'Post Month', 'Post Weekday', 'Post Hour', 'Paid', 'Total Interactions']]

# 2
df_OHE = pd.get_dummies(df_sub4, columns=['Type', 'Category', 'Post Month', 'Post Weekday', 'Post Hour'])

# 3
scaler = MinMaxScaler()
df_OHE[['Page total likes', 'Total Interactions']] = scaler.fit_transform(df_OHE[['Page total likes', 'Total Interactions']])

# 4
X = df_OHE.drop('Total Interactions', axis=1)
y = df_OHE['Total Interactions']
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=42)

# 5
model = Sequential()
model.add(Dense(64, input_dim=X_train.shape[1], activation='relu'))
model.add(Dense(32, activation='relu'))
model.add(Dense(1, activation='sigmoid'))
model.compile(optimizer='adam', loss='mean_squared_error', metrics=[metrics.MeanSquaredError()])

model.fit(X_train, y_train, epochs=150, batch_size=10, verbose=0)
mse = model.evaluate(X_test, y_test, verbose=0)[1]
mse


```

Output:
```
c:\Users\Justin\anaconda3\envs\MGMT-4190\Lib\site-packages\keras\src\layers\core\dense.py:87: UserWarning: Do not pass an `input_shape`/`input_dim` argument to a layer. When using Sequential models, prefer using an `Input(shape)` object as the first layer in the model instead.
  super().__init__(activity_regularizer=activity_regularizer, **kwargs)
0.0020866014529019594
```

<!-- grade_id: t2, points: 0, grade: true, solution: false -->
```python
#check for variable names
assert 'df_OHE' in globals()
assert 'X_train' in globals()
assert 'y_test' in globals()
assert 'mse' in globals()
assert 'model' in globals()
```

<!-- grade_id: 5, points: 8, grade: true, solution: false -->
```python
#Test cell-5 (8 points)
#DO NOT MODIFY/DELETE THIS CELL
```

<!-- grade_id: 6, points: 8, grade: true, solution: false -->
```python
#Test cell-6 (8 points)
#DO NOT MODIFY/DELETE THIS CELL
assert len(X_train)==346
assert len(y_test)==149
assert mse<0.005
assert len(model.weights)==6
```

```python

```